---
title: 'Chapter one'
description: 'It begins.'
---

## Can you flip coins?

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Welcome to a lesson on how we can generate "simulated" data that can help us understand real biological processes. 

Making up data might seem like a foreign concept, and perhaps a little like cheating, but it can be really useful. 

The logic goes like this... we create data with certain properties, or from a certain process. Then we can compare other, real data to our simulated data. If they share a lot of features, it may indicate that the real data was generated by a similar process to the one we simulated. 

Let's start with a trivial example to explain, and then we will get to an interesting biological example very soon... I promise. 

We can ask for the results of coin flips (or any other process with only two outcomes like heads/tails) with the command _rbinom_.



`@instructions`
The function _rbinom_ needs three values - the number of samples, the number of flips per sample, and the probability of a success for a flip. Let us call heads a success and tails a failure. If we wanted to flip 10 coins, each 2 times, and each was a fair coin with equal 50% probability of heads and tails, we could write `rbinom (10 ,2 , 0.5)`. How would you ask for 10 flips of each of 10 fair coins?

`@hint`
Try using the function called `rbinom`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
#flip 1 fair coin once
rbinom( , , )
```

`@solution`
```{r}
rbinom(10,10, 0.5)
```

`@sct`
```{r}
ex() %>% check_function("rbinom") %>% {
  check_arg(., "n") %>% check_equal()
  check_arg(., "size") %>% check_equal()
    check_arg(., "prob") %>% check_equal()
}
                          success_msg("Nice!")
```
