---
title: 'Chapter Two'
description: 'It continues'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: e233a3ae70
xp: 100
```

It is time again to do some coding.

`@instructions`
Use rbinom to get some expectations for wildebeests.

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
rbinom( , , )
```

`@solution`
```{r}
rbinom(1000,8, 0.8)
```

`@sct`
```{r}

```
