---
title: 'Chapter one'
description: 'It begins.'
---

## Can you flip coins?

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Welcome to a lesson on how we can generate "simulated" data that can help us understand real biological processes. 

Making up data might seem like a foreign concept, and perhaps a little like cheating, but it can be really useful. 

The logic goes like this... we create data with certain properties, or from a certain process. Then we can compare other, real data to our simulated data. If they share a lot of features, it may indicate that the real data was generated by a similar process to the one we simulated. 

Let's start with a trivial example to explain, and then we will get to an interesting biological example very soon... I promise. 

We can ask for the results of coin flips (or any other process with only two outcomes like heads/tails) with the command _rbinom_.



`@instructions`
The function _rbinom_ needs three values - the number of samples, the number of flips per sample, and the probability of a success for a flip. Let us call heads a success and tails a failure. If we wanted to flip 10 coins, each 2 times, and each was a fair coin with equal 50% probability of heads and tails, we could write `rbinom (10 ,2 , 0.5)`. How would you ask for 10 flips of each of 10 fair coins? Assign the results to a vector called _x_.

`@hint`
Try using the function called `rbinom`

`@pre_exercise_code`
```{r}
set.seed(100)
```

`@sample_code`
```{r}
#flip 10 fair coin 10 times each
x<-rbinom( , , )
```

`@solution`
```{r}
x<-rbinom(10,10, 0.5)
```

`@sct`
```{r}
ex() %>% check_object("x") %>% check_equal()
                          success_msg("Nice!")
```

---

## flipping lots of coins

```yaml
type: NormalExercise
key: 4989b6fca9
xp: 100
```

Excellent work. Now, what do those numbers mean, in your results? 

 Yep - each one is the number of successes out of 10. So, if we defined heads as success, these would be the number of heads out of 10 flips, in 10 separate samples. 

`@instructions`
So, let us pretend that we don't know anything about coins. What would happen over many, many flips? Write some code to flip one fair coin, 1000 times.

`@hint`
Try using the function called `rbinom`

`@pre_exercise_code`
```{r}
x
```

`@sample_code`
```{r}
#flip 10 fair coin 10 times each
rbinom( , , )
```

`@solution`
```{r}
rbinom(1,1000, 0.5)
```

`@sct`
```{r}
ex() %>% check_function("rbinom") %>% {
  check_arg(., "n") %>% check_equal()
  check_arg(., "size") %>% check_equal()
    check_arg(., "prob") %>% check_equal()
}
                          success_msg("Nice!")
```
